<a href="https://colab.research.google.com/github/pmSLUMSHDS/hds5210-2023/blob/main/week05/week05_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def longest(L): #defining the function for longest(L)
    """(string) -> string

    This function should find the longest string in the list of stsrings and return the position number or index of the longest string.
    The function should count the number of letters in each word in the string, going through the entire list, and returning the index of the longest string in the list.

    PSEUDOCODE:
     This line starts the count of the longest_word off at zero so that each time function loops through the list the count goes up until it is through the list
     This line starts the longest_position off at zero so that each time the function loops through the list the count will increase by one for each position until it has moved through the positions of all words in list
     starting for loop using enumnerate function.....which needs two inputs, index (x) is first, then the value (string). Enumerate function returns both the index and the value of each item as it loops though the list
     if the length of the string is greater than the longest_word....
     ....then the longest_word equals the length of the string
     ....and the longest_position will be equal to x (which is the index or position) of that longest_word
     Last line just returns the position of the longest word in the list



    >>> strings = ['happy', 'birthday', 'to', 'me']
    >>> longest(strings)
    1

    >>> strings = ['melon', 'computer', 'anniversary', 'the']
    >>> longest(strings)
    2
    """
    #Using optional book Matthes, E. Python Crash Course 3rd Edition- A Hands-On, Project-Based Introduction to Programming (2023)
    #First thoughts on this problem.....
    #Need to figure out length of wordsand position of words....and return just the position, not the length
    #Do I need to establish an empty list first?
    #Do I need to start the count at 0....or position? Or is there a way to pass through the entire list/sequence? Enumerate maybe? Check .seq, .index, enumerate()
    #Book index did not have seq or index but did have enumerate....page 331 says enumerate will return both the index and the value of each item as you loop through the entire list.....will try that...
    #Use for loop? Or if/else? Or both?

    #Was still a little confused on how to use enumerate(), Googled "how to use enumerate in for loop in python", which took me to geeksforgeeks.org where I found a section on Using Enumerate Object in Loops.
    #That finally got my for loop headed in the right direction..


    #strings = []...... on second thought don't think I need an actual list
    longest_word = 0 #started the count of the longest_word off at zero so that each time function loops through the list the count goes up until it is through the list
    longest_position = 0 #stsarted the longest_position off at zero so that each time the function loops through the list the count will increase until it has moved through the positions of all words in list

    #for string in strings:
        #if len(string) > than the next string!!!! ugh....
    for x, string in enumerate(L): #enumertate needs two inputs!! geeksforgeeks.org- Using Enumerate Object in Loops (also from book, when using enumerate, index is first, then the value. So x is my index)
        #x = string.index
        #y = len(string)
        #if len(string) > len(strings[x+1])
        if len(string) > longest_word: #if the length of the string is greater than the longest_word....
          #len(x) > len(string)
            longest_word = len(string)  #then the longest_word equals the length of the string
            longest_position = x #and the longest_position will be equal to x (which is the index or position) of that longest_word
    return longest_position #return the position of the longest word in the list


### END SOLUTION
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
Trying:
    strings = ['happy', 'birthday', 'to', 'me']
Expecting nothing
ok
Trying:
    longest(strings)
Expecting:
    1
ok
Trying:
    strings = ['melon', 'computer', 'anniversary', 'the']
Expecting nothing
ok
Trying:
    longest(strings)
Expecting:
    2
ok
```

### 25.2 Farenheit to Celsius

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION
def f_to_c(temp_in_f):  #defining the function f_to_c, meaning Farenheit to Celsius, which takes in one parameter, the temperature in degrees F, and converts to degrees C
    #This example was in this week's module!! Didn't need book :)

    """(float) -> float    inputting a float and returning a float as the output

    This function will convert a given temperature in degrees Farenheit to degrees Celsius
    Inputting a float for the temp_in_f will calculate the conversion and return the answer in temp_in_c to one decimal place.

    PSEUDOCODE:
    This is the equation to convert deg F to deg C, take the value entered for deg F and subtract 32. Then mulitply by 5 and divide that by 9. That number equals the temp_in_c
    This returns the temperature in Celsius rounded to one decimal place


    >>> temp_in_f = (212)
    >>> f_to_c(212)
    100.0

    >>> temp_in_f = (32)
    >>> f_to_c(32)
    0.0

    >>> temp_in_f = (98.6)
    >>> f_to_c(98.6)
    37.0
    """


    temp_in_c = (temp_in_f - 32) * 5/9 #take the value entered for deg F and subtract 32. Then mulitply by 5 and divide that by 9. That number equals the temp_in_c
    return round(temp_in_c, 1) #returns the temperature in Celsius rounded to one decimal place



### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    temp_in_f = (212)
Expecting nothing
ok
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    temp_in_f = (32)
Expecting nothing
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    temp_in_f = (98.6)
Expecting nothing
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates. Your function should should take three parameters: a start date, an end date, and a diagnosis.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
def los(admit, discharge, diagnosis = ""): #defining the function los which stands for length of stay and will take 3 parameters ....admit is the admission date, discharge is the discharge date and
 #diagnosis is the condition diagnosed. If the diagnosis is "Observation", the los will be zero, reagadless of other 2 date inputs.

    """(date, date, str) -> int   Inputting 2 dates and a string to get an integer as the output

    This function will calculate the length od stay in whole days, given the admit date, discharge date, and the diagnosis. If "Observation" is inputted for diagnosis, then the function will return an los of 0.

    Inputting 2 dates and a string diagnosis the function will subtract the dates and return a number in whole days. Exceptions being if both dates are the same or if "Observation" is the diagnosis, then function will
    not need to do any calculations and just return the los as 0.

    PSEUDOCODE:
    this first line deals with the exception if the "Observation" was entered for the diagnosis....
    .....then the length of stay equals 0
    code for otherwise....
    this line returns the value of the admit date subtracted from the discharge date


    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5

    >>> los(date(2020,1,5), date(2020,1,8), "Observation")
    0

    >>> los(date(2023,8,11), date(2023,8,21), "Asthma")
    10
    """

    if diagnosis == "Observation":  #this line deals with the exception if the "Observarion" was entered for the diagnosis....
      return 0   #then los is 0
    else: #otherwise....
      return (discharge - admit).days #return the value of the admit date subtracted from the discharge date



### END SOLUTION
```

```python
assert los(date(2019,1,3), date(2019,1,8), "COPD") == 5
assert los(date(2020,1,5), date(2020,1,8), "Observation") == 0
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2019,1,3), date(2019,1,8), "COPD")
Expecting:
    5
ok
Trying:
    los(date(2020,1,5), date(2020,1,8), "Observation")
Expecting:
    0
ok
Trying:
    los(date(2023,8,11), date(2023,8,21), "Asthma")
Expecting:
    10
ok
```

```python
encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for encid, visit in encounters.items():
    print(los(visit['admit'],visit['discharge'],visit['diagnosis']))

# Your output should be
# 5
# 4
# 1
# 0
# 0
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
from datetime import date
### BEGIN SOLUTION
def average_los(encounters):  #defining the average length of stay which takes as its only input a parameter that is a dictionary of dictionaries called encounters. This information is given by patient ID and gives the
#admission date, the discharge date and the diagnosis. If the length of stay of any individual encounter = 0, then that encounter will not be counted towards the average_los. And if the diagnosis is entered at "Observation",
#then the average_los will also be 0.

    """(list) -> float    Enter a list and returns a float as the output
    Function takes the encounters, a dictionarly of dictionaries, that gives the admit date, discharge date, and the diagnosis, if provided and calculates the average_los.

    Wasn't completely sure if los from 25.3 would carry over here and save a calculation. Tried to follow the example from class video, but was not working, so will add los math. Had to redo, step by step as I said it out loud :)

    Function takes in 2 dates and a diagnosis and calculates the average los for the patients given in list. Exceptions being if admit and discharge dates are the same (so los = 0) or if "Observation" is the diagnosis, both of
    those scenarios will return the los as 0 and not be counted towards the average.

    PSEUDOCODE:
    This line starts the count of the total length of stay off at zero so that each time function loops through the list the count goes up by one until it is through the list
    This line starts the total patients (or encountrs) off at zero so that each time the function loops through the list the count will increase by one for each encounter until it has moved through the list
    Starting for loop using .items function to iterate through the dictionary named encounters and in next 3 lines use the variables assigned for the dates and length of stay
    Assigned whatever date is entered for 'admit' to the variable date_admit
    Assigned whatever date is entered for the 'discharge' to the variable date_discharge
    This is calculation for obtaining the length of stay for that encounter. Subtracting the admit date from the discharge date.
    This line deals with the excpetions if the length of stay does not calculate to 0 or if observation is not entered for the diagnosis....
    .....then the total length of stay will be the length of stay plus one to keep a running count of total days
    .....and the number of patients (or encounters) also increases by one to keep a running count of the encounters that will count towars the average
    if the total number of patients (or encounters) is greater than 0....
    .....then the average length of stay is equal to the total length of stay divided by the total number of patients (or encounters)
    otherwise.....
    .....the average length of stay equals 0
    return the average length of stay

    >>> round(average_los(encounters), 2)
    3.33
    """

    total_los = 0
    #total_encounters = len(enounters) tried to use same format as video in this week's module, but can't make it work.....strange....rethinking this....
    total_patients = 0   #changed this to patients as encounters was appearing way too many times and started confusing me

    #for encounter in encounters:
    for id, encounter in encounters.items(): #we used this format last week for the families list....tried it here with the .items
        #enc_los = los(enc['admit'], enc['discharge'], enc['diagnosis'])
        date_admit = encounter['admit'] #separated out the individual parameters, so I could make sense of it, assigned whatever date is entered for 'admit' to the variable date_admit
        date_discharge = encounter['discharge'] #assigned whatever date is entered for the 'discharge' to the variable date_discharge
        los = (date_discharge - date_admit).days #pulled the .days from previous problem, this is calculation for obtaining the length of stay for that encounter. Subtracting the admit date from the discharge date.
        #if enc_los != 0
        if los != 0 and encounter['diagnosis'] != "Observation": #this line deals with the excpetions if the length of stay calculates is not 0 or if observation is not entered for the diagnosis....
              #total_los += enc_los
              total_los += los #then the total length of stay will be the length of stay plus one
              total_patients += 1 #and the number of patients also increases by one

    if total_patients > 0: #if the total number of patients (or encounters) is greater than 0....
        average_los = total_los / total_patients #then the average length of stay is equal to the total length of stay divided by the total number of patients (or encounters)
    else: #otherwise.....
        average_los = 0 #the average length of stay equals 0

        #else:
            #total_encounters -= 1

    #if total_encounters == 0:
        #return 0

    return average_los #returns the average length of stay

### END SOLUTION
```

```python
assert (round(average_los(encounters),2)) == 3.33
```

```python

encounters = {
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(round(average_los(encounters),2) == 3.33)
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    round(average_los(encounters), 2)
Expecting:
    3.33
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
def c_to_f(temp_in_c):  #defining a function clled c_to_f that takes one input, the temperature in Celsius using equation temp_in_f = (temp_in_f * 9/5) + 32
    #This problem is technicallly the revers of the previous problem above, so no book or other references needed.

    """(float) -> float    inputting a float and returning a float as the output

    This function will convert a given temperature in degrees Celsius to degrees Farenheit
    Inputting a float for the temp_in_c will calculate the conversion and return the answer in temp_in_f to one decimal place.

    PSEUDOCODE:
    This line calculates the deg F. Usng the value given for deg C, you multiply by 9 and then divide by 5 and then add 32. The resulting number is the temp_in_f.
    This line returns the temperature in degrees Farenheit rounded to one decimal place.



    >>> temp_in_c = (100)
    >>> c_to_f(100)
    212.0

    >>> temp_in_c = (0)
    >>> c_to_f(0)
    32.0

    >>> temp_in_c = (37)
    >>> c_to_f(37)
    98.6

    >>> temp_in_c = (56)
    >>> c_to_f(56)
    132.8
    """


    temp_in_f = (temp_in_c * 9/5) + 32 #this line calculates the deg F. Usng the value given for deg C, you multiply by 9 and then divide by 5 and then add 32. The resulting number is the temp_in_f
    return round(temp_in_f, 1) #returns the temperature in degrees Farenheit rounded to one decimal place


### END SOLUTION
```

```python
assert c_to_f(100) == 212.0
assert c_to_f(0) == 32.0
assert c_to_f(37) == 98.6
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    temp_in_c = (100)
Expecting nothing
ok
Trying:
    c_to_f(100)
Expecting:
    212.0
ok
Trying:
    temp_in_c = (0)
Expecting nothing
ok
Trying:
    c_to_f(0)
Expecting:
    32.0
ok
Trying:
    temp_in_c = (37)
Expecting nothing
ok
Trying:
    c_to_f(37)
Expecting:
    98.6
ok
Trying:
    temp_in_c = (56)
Expecting nothing
ok
Trying:
    c_to_f(56)
Expecting:
    132.8
ok
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week03 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -a -m "message"`)
  5. Push your changes (`git push`)
  
If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.