<a href="https://colab.research.google.com/github/pmSLUMSHDS/hds5210-2023/blob/main/week04/week04_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 4 Exercises

See: _McKinney 3.1_ and [Python Documentation](https://docs.python.org/3/tutorial/datastructures.html) section 5 on data structures.

Modules 13 through 15 have all been about different built-in data types that can be used to create larger and more complex data structures.  We learned about tuples, lists, dictionaries, and how to access and process that data in new ways.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 20.1 Create a sorted list of names

Take the list of names below (some are first and last, others are first, middle, and last) and produce a final list of full names.

```
sort_names([
  ['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']])

['Paul Boal', 'Donald Duck', 'Kermit the Frog']
```

**Tip:** If you have a list of lists and want to sort it, Python can do that. It will simply sort using the first item in each of the lists inside your list.

```python
names = [['Paul', 'Boal'], # this is list of names...or nested list of nmes....to be used in the code
  ['Kermit','the','Frog'],
  ['Donald','Duck']]
```

```python

```

I am unsure why this was soooo much harder than the math section from last week. I used the optional textbook (Matthes, E. Python Crash Course 3rd Edition- A Hands-On, Project-Based Introduction to Programming, 2023). and watched a few YouTube videos (https://www.youtube.com/watch?v=3Tcq_dBhIrQ, https://www.youtube.com/watch?v=1lnPvBUSseA&t=98s) on lists and python for beginners. But they only served to confuse me more. So I started moving things around as you can see from all the lines turned into comments. Still didn't work.

```python
from typing_extensions import LiteralString #importing necessary functions
def name_sorter(names_list): #def is short for define....so we are defining the function, calling it name_sorter. Inside the () is what will be inputted to make the function run.
    # Do something
    full_names = [] #naming the list
    for name in names_list: #starting the loop, for each name in the names_list....do something...in the next line
        full_names.sort #(key = lambda k: k.split()[-1])      #sorting the list
        #full_names.append(name[0] + " " + name[-1])
        #full_names.append(name.split())
    #names_list = []
    for name in sorted(full_names, key=lambda x: x[-1]):   #starting another loop....for each name in the new full_names list, which is sorted first....not completely sure why we called lambda here, it is some sort of temporary function and the remainder of the lines relates to it.
        names_list.append(' '.join(name)) #in short it loops through names keeping the names in each list together (via join), so it doesn't matter if there is a middle name or not.
        #full_names.sort(names[-1])
        #if name in names_list == (["Kermit"]): #I've tried checking the length of words or if it startswith K or if the first name is Kermit.....nothing works. But if I take out the == Kermit bit it will add 'the' to the first name in the next line. So the error has to be in this line.
            #full_names.append([name[0] + " the"])
            #full_names.sort()
        #else:
            #full_names.append(name[0] + ' ' + name[-1])
            #full_names.sort
    #full_names.append(name[0] + " " + name[-1])
    #full_names.sort
    return names_list #after looping through names in list of lists, a single list is returned sorted alphabetically by last name but with the first name listed first.
```

Okay, after the flurry of emails regarding this assignment, I Googled "how to sort a nested list of names when some have included a middle name in python". That led me to geeksforgeeks.org https://www.geeksforgeeks.org/python-program-to-sort-a-list-of-names-by-last-name/ which had several suggestions. Method 2 was not happy with me, so I switched to Method 1, which didn't work either but I tweaked it here and there. I can see now that I was making my code too specific in my if loop, so code would not have worked for any names given, only the original list, had I been able to get it to work. This is still not accurate but a little better as now it will work if middle names are included. Order still not right and it is a list of lists instead of just one list......alas......

```python
name_sorter(names)
```

Output:
```
[['Paul', 'Boal'], ['Kermit', 'the', 'Frog'], ['Donald', 'Duck']]
```

```python
names = [['Paul', 'Boal'],
  ['Kermit','the','Frog'],
  ['Donald','Duck']]

assert name_sorter(names) == ["Paul Boal", "Donald Duck", "Kermit the Frog"]
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-9-fe087893c352> in <cell line: 5>()
      3   ['Donald','Duck']]
      4 
----> 5 assert name_sorter(names) == ["Paul Boal", "Donald Duck", "Kermit the Frog"]

AssertionError:
```

### 20.2 Convert a dictionary to a list

In this exercise, we're going to take a dictionary that has families and their names, and convert that into a list of names.  The dictionary is structured such that the key is the family name, and the value for that key is a list of first names.

```
{
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Elliie'],
 ['Lester-Boal', 'Ada'],
 ...
```

I watched several videos https://www.youtube.com/watch?v=oVLQTyvu4lo, https://www.youtube.com/watch?v=1lnPvBUSseA&t=98s, which did not help me so I went back to the book (Matthes, E. Python Crash Course 3rd Edition- A Hands-On, Project-Based Introduction to Programming, 2023) and reread the dictionary and lists sections but it was actually the csdojo videos linked in the modules that got me going in the right direction. He did one with key, values in d.items that helped me kind of reason through. Not really sure why it works or if it is truly correct. Why did I need brackets and parentheses?

```python
from os import listdir #importing functions needed
def dict_to_list(families): #defining a new function called dict_to_list. Starting with a dictionary and turning to list. Families in () is what will be inputted to run code
    # Do something
    list_of_names = [] #naming new list as list_of_names
    for key, values in families.items(): #starting loop that for each set of key, values in the families
        #for key in families.items:
        for value in values: #starting another loop that for each value in values...
            #list_of_names.append(key)
            #list_of_names.append([key])
            list_of_names.append([key, value]) #take the matching key/values and append to the list_of_names
    return list_of_names #return thenewly created list_of_names
```

```python
families = {                           #this is the dictionary of names to turn into the list
  'Boal': ['Paul', 'Anny', 'Jim'],
  'Lester-Boal': ['Ellie', 'Ada', 'Teddy'],
  'Lester': ['Sarahlynn', 'Jessica', 'Grace', 'Carolynn']
}

dict_to_list({'Duck': ['Hewey','Dewey','Louie'],
              'Mouse': ['Mickey','Minney']})
```

Output:
```
[['Duck', 'Hewey'],
 ['Duck', 'Dewey'],
 ['Duck', 'Louie'],
 ['Mouse', 'Mickey'],
 ['Mouse', 'Minney']]
```

```python
dict_to_list(families)
```

Output:
```
[['Boal', 'Paul'],
 ['Boal', 'Anny'],
 ['Boal', 'Jim'],
 ['Lester-Boal', 'Ellie'],
 ['Lester-Boal', 'Ada'],
 ['Lester-Boal', 'Teddy'],
 ['Lester', 'Sarahlynn'],
 ['Lester', 'Jessica'],
 ['Lester', 'Grace'],
 ['Lester', 'Carolynn']]
```

### 20.3 Convert a list to a dictionary

Take the list of lists below and convert it into a dictionary where the entries in the dictionary are keyed off of whatever values are in column 0, and the dictionary values are the list of all the values that appear with that key from the input list.

Just tried to do the opposite of 20.2 once I got that one to work.....again, I'm not sure how or why. How can we call the actual list of names "names" and also have 'names' in list_to_dict, yet they aren't the same. Or are they? Then we can say name in names in the for loop.....how does it know what name is vs names? I feel like I keep using the same word over and over for different things so I'm never sure what is what. I am figuring these out purely by placement....which word is in parentheses and how many places do I see that word. And then which word is used for the list and where do I see that word placed, etc. If it doesn't run, start swapping words :) This plan will not get me very far.....

<!-- grade_id: average-answer, points: 0, grade: false, solution: true -->
```python
def list_to_dict(names): #defining a new function called list_to_dict. Starting with a list and turning into a dictionary using (names)

    names = [['Boal', 'Paul'],  #list of lists to be inputted and made into dictionary
             ['Duck', 'Donald'],
             ['Duck', 'Daisy'],
             ['Boal', 'Ada'],
             ['Boal', 'Teddy'],
             ['Westhus', 'Eric']]

    # To help you, I've provided a blank dictionary named `families`.

    families = {} #new dictinary called families

    # Your code should loop through `names` and do something to build
    # up `families` to be the output dictionary shown in the test
    # cells below.

    # Put your solution here
    for name in names: #for each name in names list...do something....
        key, value = name #need both the first name and last name included....so each key (last name) and value (first name)
        #if key in names: #need to have key (or last name fixed) and throw all first names in the []
        if key in families: #names didn't work, swithced to families....If the key or last name in the dictionary....
          #families.append(value)
            families[key].append(value) #....is the same then append the value and add the first name to the end of the list
        else:
          #families = value
            families[key] = [value] #keep the key/vales pairs together
    return families #return the newly created dictionary of names

```

```python
list_to_dict(names)
```

Output:
```
{'Boal': ['Paul', 'Ada', 'Teddy'],
 'Duck': ['Donald', 'Daisy'],
 'Westhus': ['Eric']}
```

<!-- grade_id: average-tests, points: 5, grade: true, solution: false -->
```python
answer = {
 'Boal':    ['Paul', 'Ada', 'Teddy'],
 'Duck':    ['Donald', 'Daisy'],
 'Westhus': ['Eric']
}

assert (list_to_dict(names) == answer)
```

### 20.4 Join using a dictionary

We have a list of patients, diagnosis, and length of stay.  We also have a dictionary that contains diagnosis and average length of stay.  Produce an output list that lists the patient and an indicator if the patient's stay was 'too long', 'too short', 'just right'

<!-- grade_id: cell-f3813e33810dc898, points: 0, grade: false, solution: true -->
```python
def patient_los(avg_los, patients):  #defining a new function called patient_los (patient length of stay) where the avgerage lengthof stay (avg_los) and patients will be inputted

    conclusions = [] #newly created list called conclusions
    # Put your solution here
    for values in avg_los: #starting loop that for every number or value in avg_los...do something....
        if avg_los < values(avg_los): #if number entered for avg_los is less than the number in the dictionary's value for avg_los, then print 'too short'
            print('too short')
        elif avg_los == values(avg_los): #else if the number entered for avg_los is equal to the number in the dictionary's value for avg_los, then print 'just right'
            print('just right')
        else:       #else if number entered for avg_los is anything else (meaning greater than) the number in the dictionary's value for avg_los, then print 'too long'
            print('too long')
    return conclusions #return the conclusions
```

this problem has given me a headache :) I will be reading the book like its a bedtime story this week. I am enjoying W3resource and geeksforgeeks....lots of examples to work through, which I clearly need.

```python
avg_los = {
    "Hemolytic jaundice and perinatal jaundice" : 2,
    "Medical examination/evaluation" : 3.2,
    "Liveborn" : 3.2,
    "Trauma to perineum and vulva" : 2.1,
    "Normal pregnancy and/or delivery" : 2,
    "Umbilical cord complication" : 2.1,
    "Forceps delivery" : 2.2,
    "Administrative/social admission" : 4.2,
    "Prolonged pregnancy" : 2.4,
    "Other complications of pregnancy" : 2.5
}

patients = [
    ['Boal', 'Medical examination/evaluation', 1.1],
    ['Boal', 'Other complications of pregnancy', 3.3],
    ['Jones', 'Liveborn', 3.2],
    ['Ashbury', 'Forceps delivery', 2.0]
]

patient_los(avg_los, patients)
```

Output:
```
<ipython-input-176-7f536fc3689c> in <cell line: 21>()
     19 ]
     20 
---> 21 patient_los(avg_los, patients)

<ipython-input-175-13c9373f6d33> in patient_los(avg_los, patients)
      4     # Put your solution here
      5     for values in avg_los:
----> 6         if avg_los < values(avg_los): # need to say if # entered for avg_los is less than the number in the dictionary for avg_los.....not sure how to code that if both are called avg_los
      7             print('too short')
      8         elif avg_los == values(avg_los):

TypeError: 'str' object is not callable
```

```python
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Mounted at /content/drive
```

<!-- grade_id: cell-2af0a520bb56ee21, points: 10, grade: true, solution: false -->
```python
conclusions = [['Boal', 'too short'],
    ['Boal', 'too long'],
    ['Jones', 'just right'],
    ['Ashbury', 'too short']]

assert(patient_los(avg_los, patients) == conclusions)
```

```python

```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week04/week04_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

