<a href="https://colab.research.google.com/github/pmSLUMSHDS/hds5210-2023/blob/main/week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="allergies.json"
```

Purpose of function & what it will do:

The purpose of this function is to count the number of entries or items in the dictionary of lists of dictionaries.

Defined a function called allergy_count that has ALLERGIES_FILE as the only input. The input file is a string.

The function will evaluate the file, count the separate entries and return the output as an integer.

PSEUDOCODE:

Define a function called allergy_count that takes as its only input the JSON file, ALLERGIES_FILE.

Open the JSON file as the variable file, load the file, and save it as medical_file.

Variable entries will be equal to the medical_file list, entry.

Variable number_items is equal to the length of the variable entries.

Return the number_items.

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python

# Put your solution here
def allergy_count(ALLERGIES_FILE):

    """(string) -> int
    Function evaluates the file, counts items, and returns an integer

    >>> allergy_count('allergies.json')
    4
    """


    with open("allergies.json") as file:
        medical_file = json.load(file)
    #number_items = 0
    entries = medical_file['entry']
    number_items = len(entries)

    return number_items
#count = allergy_count("allergies.json")
#print(count)



```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
import doctest
doctest.run_docstring_examples(allergy_count, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_count('allergies.json')
Expecting:
    4
ok
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

Purpose of function & what it will do:

The purpose of this function is to count the number of unique patient names in the dictionary of lists of dictionaries.

Defined a function called patient_count that has ALLERGIES_FILE as the only input. The input file is a string.

The function will evaluate the file, count the number of different patients and return the output as an integer.

PSEUDOCODE:

Define a function called patient_count that takes as its only input the JSON file, ALLERGIES_FILE.

Open the JSON file as the variable file, load the file, and save it as medical_file.

Empty variable patient_set will be used to hold the unique names, removing duplicates.

For each entry in the medical_file named entry, need to drill down to what specifically we are trying to count.

Inside the dictionary of lists named entry, find resource, then patient, then display and save that string as patient_name.

Append patient_set with each different name found.

Variable number_names is equal to the length of (or number of names in) the patient_set.

Return the number_names.

```python
# Put your solution here
def patient_count(ALLERGIES_FILE):

    """(string) -> int
    Function evaluates the file, counts unique patient names, and returns an integer

    >>> patient_count('allergies.json')
    2
    """

    with open("allergies.json") as file:
        medical_file = json.load(file)

    #patients = medical_file['entry']
    #number_patients = 1
    #for key, value in medical_file.items:
        #key : value = "patient : display"
    #if value == value:
        #number_patients == number_patients
    #else:
        #number_patients += 1
    #return number_patients

#First attmept did not work so I Googled "how to count names in a file" which took me to geeksforgeeks.org where they explained that the set() function will automatically remove duplicates while counting unique names.
#Then you just need to add it to the set().
#In trying to see how to drill down to specific info in JSON file, pg 346 in the optional class book, Matthes, E."Python Crash Course- A Hands-On, Project-Based Introduction to Programming" (2023)
#There's an example at the botton of the page that used brackets to get to the info needed, so used brackets to get down to the name. Also saw the brackets in the weekly modules for JSON.

    patient_set = set()
    for entry in medical_file['entry']:
        patient_name = entry['resource']['patient']['display']
        patient_set.add(patient_name)

    number_names = len(patient_set)

    return number_names

```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

```python
import doctest
doctest.run_docstring_examples(patient_count, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    patient_count('allergies.json')
Expecting:
    2
ok
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


Purpose of function & what it will do:

The purpose of this function is to count the number of allergies per patient names in the dictionary of lists of dictionaries.

Defined a function called allergy_per_patient that has a JSON file, ALLERGIES_FILE, as the only input. The input file is a string.

The function will evaluate the file, find the different patients and allergies per patient and return the output as string : integer (patient name: # of allergies).

PSEUDOCODE:

Define a function called allergy_per_patient that takes as its only input the JSON file, ALLERGIES_FILE.

Open the JSON file as the variable file, load the file, and save it as medical_file.

Established empty dictionary allergy_number{}.

For each entry in the medical_file, need to drill down to what specifically we are trying to count.

Inside the dictionary of lists named entry, find resource, then patient, then display to locate the patient's name and save to variable patient_name.

If the name in the dictiontary allergy_number[name] is the same, then add one.

Otherwise keep allergy_number[name] total at 1.

Return the allergy_number count.

Established another variable, allergies, for the store the dictionary and used print(f"...") to print the patient name and number of allergies for that patient.


```python
# Put your solution here
def allergy_per_patient(ALLERGIES_FILE):

    """(string) -> string : int
    Function evaluates the file, counts number of allergies for each patient, and returns the patient name and number of allergies per patient in a dictionary
    """

    with open("allergies.json") as file:
        medical_file = json.load(file)

    """patient_set = set()
    allergy_set = set()
    for entry in medical_file['entry']:
        patient_name = entry['resource']['patient']['display']
        patient_set.add(patient_name)
    for entry in medical_file['entry']:
        allergy = entry['resource']['substance']['text']
        allergy_set.add(patient_name)
    allergy_number = len(allergy_set)
    return patient_set, allergy_number
    """
# first try was super close, tried to use set() again, but didn't return the number of allergies for both names, trying again with an empty dictionary instead of set()...

    allergy_number = {}

    for entry in medical_file['entry']:
         patient_name = entry['resource']['patient']['display']
         if patient_name in allergy_number:
             allergy_number[patient_name] += 1
         else:
             allergy_number[patient_name] = 1

    return allergy_number
    #print(f"{name}: allergy_count")- had to Google "how to use f in print()" to figure out how to print the answer in proper format which took me back to geeksforgeeks.org. Read section on
    #f-strings in Python and also stackoverflow.com had a section on What is print(f"...") that was also helpful.

allergies = allergy_per_patient("allergies.json")
for patient_name, count in allergies.items():
    print(f"{patient_name}: {count}")

   # answer is not in the same order, but not sure how to switch it....opposite of alphabetical maybe. Or maybe least amount of allergies to highest....?


```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

Purpose of function & what it will do:

The purpose of this function is to list each patient as presented in the file, along with their allergy and reaction to that allergy.

Defined a function called allergy_list that has a JSON file, ALLERGIES_FILE, as the only input. The input file is a string.

The function will evaluate the file, find each patient entry, the allergy for that entry, and the patient's reaction to the allregy listed and return the output as a list per instance of the patient name, name of allergy, and allergy reaction.

PSEUDOCODE:

Define a function called allergy_list that takes as its only input the JSON file, ALLERGIES_FILE.

Open the JSON file as the variable file, load the file, and save it as medical_file.

Established empty list named patient_list.

For each patient in the medical_file, need to drill down to what specifically we are trying to find.

Inside the dictionary of lists named entry, find resource, then patient, then display to locate the patient's name and save to variable patient_name.

Inside the dictionary of lists named entry, find resource, then substance, then text to locate the patient's allergy and save to variable allergy_type.

Inside the dictionary of lists named entry, find resource, then reaction, then manifestation, then text to locate the patient's allergy reaction and save to variable rxn.

Append all three to the list patient_list[].

Return the patient_list.


<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_list(ALLERGIES_FILE):

    """(string) -> [string, string, string] as name, allergy, reaction
    Function evaluates the file, finds each patient entry, and returns the patient name, the allergy,
    and the reaction.
    """
    with open("allergies.json") as file:
        medical_file = json.load(file)

    patient_list = []

    for entry in medical_file['entry']:
        patient_name = entry['resource']['patient']['display']
        allergy_type = entry['resource']['substance']['text']
        rxn = entry['resource']['reaction'][0]['manifestation'][0]['text'] #couldn't just list these like the others, per the instructions the list of lists changed things. Googled "how to select the correct item in a json file" which took me to stackoverflow.com which had an example titled "How to access first element of JSON object array" which used [0] to choose correct position.
        patient_list.append([patient_name, allergy_type, rxn])

#patients = allergy_list("allergies.json")
#for patient in patients:
    #for reaction in patients:
        #manifest = reaction['manifestation'][0]['text']
    #print(f"{patient_name}, {allergy_type}, {manifestation}")
    #print([patient[0], patient[1], 'rxn'])

    return patient_list
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.


Purpose of function & what it will do:

The purpose of this function is to return the patient's reaction to the allergry when given the patient's name, allergy and the file as inputs.

Defined a function called allergy_reaction that has a JSON file, ALLERGIES_FILE, the patients name, and the substance they are allergic to as the 3 string inputs.

The function will evaluate the file, patient name and substance and return the output as the patient's reaction to that allergy as a string.

PSEUDOCODE:


Define a function called allergy_reaction that takes as its input the JSON file, patient's name, and allergic substance.

Open the JSON file as the variable file, load the file, and save it as medical_file.

Established empty list named patient_rxns.

For each patient in the medical_file, need to drill down to what specifically we are trying to find.

Inside the dictionary of lists named entry, find resource, then patient, then display to locate the patient's name and save to variable patient_name.

Inside the dictionary of lists named entry, find resource, then substance, then text to locate the patient's allergy and save to variable allergy_type.

Inside the dictionary of lists named entry, find resource, then reaction, then manifestation, then text to locate the patient's allergy reaction and save to variable rxn.

If the substance is in the medical_file, return the patient's reaction to it.

If the substance is not in the medical_file, return "None".



<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here

def allergy_reaction(ALLERGIES_FILE, patient, substance):
#def allergy_reaction(ALLERGIES_FILE, allergy_list(ALLERGIES_FILE)) #could not figure out calling a function in a function
#def allergy_reaction(ALLERGIES_FILE, patient_list):

    """(string, string) -> (string)
    Function evaluates the file, and 2 other inputs of patient name and allergy and returns the patient reaction
    to that allergy.
    """
    with open(ALLERGIES_FILE) as file:
        medical_file = json.load(file)

    #patient_rxns = []
    #med_file = allergy_list.items

    #react = allergy_reaction(ALLERGIES_FILE, patient, substance)
    #for patient_name, substance in patient_list.items():

    for entry in medical_file['entry']:
        patient_name = entry['resource']['patient']['display']
        #substance = entry['resource']['substance']['text']
        rxn = entry['resource']['reaction'][0]['manifestation'][0]['text']
        for substance in medical_file['entry']:
              #print(rxn)
            if substance not in medical_file['entry']:
                rxn == 'None'
                print(rxn)
            else:
                return rxn
            #substance = None
            #patient_rxns.append([patient_name, substance, rxn])
            #return rxn
            #print(rxn)
        #else:
        #for substance not
            #substance not in medical_file['entry']
            #patient_rxns.append([patient_name, substance, rxn])
            #rxn == 'None'

            #return rxn

        #return patient_rxns
        #print(f"(ALLERGIES_FILE, {patient_name}, {allergy_type}) == {rxn}")
            #return rxn
        #I Googled "how to use not in a for loop" which took me to stackoverflow.com where I read a section titled "Using 'not' in a for-loop statement".....clearly didn't help me :)
        #Tried Googling 'how to say if item is in a list in python' which took me to geeksforgeeks.org where I read about 'Check if element exists in list in python' but that did not help either.



    #thought I could just print differently. But quotes are not correctly placed like the answer, which likely means I need to redo code to get the proper print result as opposed to
    #just having it print in the answer in the right format.

#edited code and now my loop is not looping....not sure how to correct this....sure it's an easy fix, but don't see it....apologies....
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-250-2a5234be0de8> in <cell line: 2>()
      1 assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
----> 2 assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
      3 assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
      4 assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
      5 assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'

AssertionError:
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

